<a href="https://colab.research.google.com/github/TongSii/hds5210-2025/blob/main/week05/week05_assignment_coding.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
def longest(L):
    """
    Return the position of the longest string in L.

    >>> longest(['happy', 'birthday', 'to', 'me'])
    1
    >>> longest("enjoy class".split(" "))
    0
    >>> longest(['when','what','where','how','who'])
    2
    """
    # find the longest string
    longest_string = max(L, key=len)
    # return its position
    return L.index(longest_string)
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    longest(['happy', 'birthday', 'to', 'me'])
Expecting:
    1
ok
Trying:
    longest("enjoy class".split(" "))
Expecting:
    0
ok
Trying:
    longest(['when','what','where','how','who'])
Expecting:
    2
ok
```

### 25.2 Farenheit to Celsius

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
def f_to_c(fahrenheit):
    """
    Convert degrees Fahrenheit to degrees Celsius.

    >>> f_to_c(212)
    100.0
    >>> f_to_c(32)
    0.0
    >>> f_to_c(98.6)
    37.0
    """
    return 5/9 * (fahrenheit - 32)
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    f_to_c(212)
Expecting:
    100.0
ok
Trying:
    f_to_c(32)
Expecting:
    0.0
ok
Trying:
    f_to_c(98.6)
Expecting:
    37.0
ok
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates. Your function should should take three parameters: a start date, an end date, and a diagnosis.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

### BEGIN SOLUTION
from datetime import date

def los(start_date, end_date, diagnosis):
    """
    Calculate the Length of Stay (LOS) for a patient,
    if the diagnosis is "Observation" then the LOS is 0,
    otherwise it is the difference in days between the start and end dates.

    >>> los(date(2019,1,3), date(2019,1,8), "COPD")
    5
    >>> los(date(2020,1,5), date(2020,1,8), "Observation")
    0
    """
    if diagnosis == "Observation":
        return 0
    return (end_date - start_date).days

encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

for encid, visit in encounters.items():
    print(los(visit['admit'],visit['discharge'],visit['diagnosis']))

# Doctests
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)

### END SOLUTION
```

Output:
```
5
4
1
0
0
Finding tests in NoName
Trying:
    los(date(2019,1,3), date(2019,1,8), "COPD")
Expecting:
    5
ok
Trying:
    los(date(2020,1,5), date(2020,1,8), "Observation")
Expecting:
    0
ok
```

```python
assert los(date(2019,1,3), date(2019,1,8), "COPD") == 5
assert los(date(2020,1,5), date(2020,1,8), "Observation") == 0
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    los(date(2019,1,3), date(2019,1,8), "COPD")
Expecting:
    5
ok
Trying:
    los(date(2020,1,5), date(2020,1,8), "Observation")
Expecting:
    0
ok
```

```python
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for encid, visit in encounters.items():
    print(los(visit['admit'],visit['discharge'],visit['diagnosis']))

# Your output should be
# 5
# 4
# 1
# 0
# 0
```

Output:
```
5
4
1
0
0
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python

### BEGIN SOLUTION
from datetime import date

def los(start_date, end_date, diagnosis):
    """
    Calculate the Length of Stay (LOS) for a patient,
    if the diagnosis is "Observation" then the LOS is 0,
    otherwise it is the difference in days between the start and end dates.
    """
    if diagnosis == "Observation":
        return 0
    return (end_date - start_date).days

def average_los(encounters):
    """
    Calculate the average Length of Stay (LOS), not considering individual encounters
    with a LOS of 0.

    >>> round(average_los(encounters),2)
    3.33
    """
    total_days = 0
    total_encounters = 0

    for encid, visit in encounters.items():
        los_value = los(visit['admit'],visit['discharge'],visit['diagnosis'])
        if los_value > 0:
            total_days += los_value
            total_encounters += 1

    return total_days / total_encounters if total_encounters else 0

# Given Encounters data
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

# Perform Assertion Test
assert(round(average_los(encounters), 2) == 3.33)

# Run Doctests
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)


### END SOLUTION
```

Output:
```
Finding tests in NoName
Trying:
    round(average_los(encounters),2)
Expecting:
    3.33
ok
```

```python
from datetime import date
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(round(average_los(encounters),2) == 3.33)
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    round(average_los(encounters),2)
Expecting:
    3.33
ok
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
def c_to_f(celsius):
    """
    Convert degrees Celsius to degrees Fahrenheit.

    >>> c_to_f(100)
    212.0
    >>> c_to_f(0)
    32.0
    >>> c_to_f(37)
    98.6
    """
    return (celsius * 9/5) + 32
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    c_to_f(100)
Expecting:
    212.0
ok
Trying:
    c_to_f(0)
Expecting:
    32.0
ok
Trying:
    c_to_f(37)
Expecting:
    98.6
ok
```

```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

Save you work to the path week05/week05_assignment_coding.ipynb

```python

```